Updating the sulcal landscape of the human lateral parieto-occipital junction provides anatomical, functional, and cognitive insights

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #2 (Public Review):

Strengths

(1) The definition of highly variable yet highly reproducible sulci such as the slocs-v feeds the community with new anatomo-functional landmarks (which is emphasized by the provision of a probability map in supp. mat., which in my opinion should be proposed in the main body).

We agree with Reviewer 2 that there is merit to including the probability maps as a main text Figure rather than Supplementary Figure. We have now added it to the main text.

Weaknesses

(1) While the identification of the sulci has been done thoroughly with expert validation, the sulci have not been labeled in a way that enables the demonstration of the reproducibility of the labeling.

Our group was unable to use an approach amenable to calculating inter-rater agreements to expedite the process of defining thousands of sulci at the individual level in multiple regions as this was our first study comprehensively documenting the sulcal organization of this region. Nevertheless, our method followed a rigorous, three-tiered procedure to ensure accurate sulcal definitions were identified in all participants. In the case of this study, authors YT and TG first defined sulci. These sulci were then checked by a trained expert (EHW). Finally, sulcal definitions were finalized by the senior author, an expert neuroanatomist (KSW). We emphasize that this process has produced reproducible anatomical results when charting other regions such as posteromedial cortex (Willbrand et al., 2023 Science Advances; Willbrand et al., 2023 Communications Biology; Maboudian et al., 2024 The Journal of Neuroscience; Ramos Benitez et al., 2024 Neuropsychologia), ventral temporal cortex (Miller et al., 2020 Scientific Reports; Parker et al., 2023 Brain Structure and Function), and lateral prefrontal cortex (Miller et al., 2021 The Journal of Neuroscience; Voorhies et al., 2021 Nature Communications; Yao et al., 2022 Cerebral Cortex; Willbrand et al., 2022 Brain Structure and Function; Willbrand et al., 2023 The Journal of Neuroscience; Willbrand et al., 2024 Brain Structure and Function) across age groups, species, and clinical populations. For the present study, by the time the final tier of our method was reached, we emphasize that a very small percentage (~2%) of sulcal definitions were actually modified. We will include an exact percentage in future publications in LPC/LOPJ.

Our Methods have been edited to describe these features (Pages 21-22):

“As this is the first time the sulcal expanse of LPC/LOPJ was comprehensively charted with a focus on pTS, the location of each sulcus was confirmed through a three-tiered procedure for each participant in each hemisphere. First, trained independent raters (Y.T. and T.G.) identified sulci. Second, these definitions were checked by a trained expert (E.H.W.). Third, these labels were finalized by a neuroanatomist (K.S.W.). We emphasize that this procedure has produced reproducible results in our prior work across the cortex (Miller et al. 2021; Voorhies et al. 2021; Yao et al. 2022; Willbrand et al. 2023; Willbrand et al. 2022; Willbrand et al. 2024; Parker et al. 2023; Miller et al. 2020; Willbrand et al. 2022; Willbrand et al. 2023; Maboudian et al. 2024; Ramos Benitez et al. 2024). All LPC sulci were then manually defined and saved as .label files in FreeSurfer using tksurfer tools, from which morphological and anatomical features were extracted. We defined LPC/LPOJ sulci for each participant based on the most recent schematics of sulcal patterning by Petrides (2019) as well as pial, inflated, and smoothed white matter (smoothwm) FreeSurfer cortical surface reconstructions of each individual. In some cases, the precise start or end point of a sulcus can be difficult to determine on a surface (Borne et al., 2020); however, examining consensus across multiple surfaces allowed us to clearly determine each sulcal boundary in each individual. For four example hemispheres with these 13-17 sulci identified, see Fig. 1a (Supplementary Fig. 5 for all hemispheres). The specific criteria to identify the slocs and pAngs are outlined in Fig. 1b.”

Reviewer #3 (Public Review):

Weaknesses

(1) The numbers of subjects are inherently limited both in number as well as in being typically developing young adults.

First, although the sample size of the present study is small in number in comparison to large N, group-level neuroimaging analyses, it is comparable to precision neuroimaging studies examining sulcal features in individual participants (for example, Cachia et al., 2021 Frontiers in Neuroanatomy; Garrison et al., 2015 Nature Communications; Lopez-Persem et al., 2019 The Journal of Neuroscience; Miller et al., 2021 The Journal of Neuroscience; Roell et al., 2021 Developmental Cognitive Neuroscience; Voorhies et al., 2021 Nature Communications; Weiner, 2019 The Anatomical Record; Willbrand, et al., 2022 Science Advances; Willbrand, et al., 2022 Brain Structure & Function; Yao et al., 2022 Cerebral Cortex). We discuss this point in detail in the Limitations subsection of the Discussion (Page 17):

“This manual method is also arduous and time-consuming, which, on the one hand, limits the sample size in terms of number of participants, while on the other, results in thousands of precisely defined sulci. This push-pull relationship reflects a broader conversation in the human brain mapping and cognitive neuroscience fields between a balance of large N studies and “precision imaging” studies in individual participants (Gratton et al., 2022; Naselaris et al., 2021; Rosenberg and Finn, 2022). Though our sample size is comparable to other studies that produced reliable results relating sulcal morphology to brain function and cognition (for example, Cachia et al., 2021; Garrison et al., 2015; Lopez-Persem et al., 2019; Miller et al., 2021; Roell et al., 2021; Voorhies et al., 2021; Weiner, 2019; Willbrand et al., 2022a, 2022b; Yao et al., 2022), ongoing work that uses deep learning algorithms to automatically define sulci should result in much larger sample sizes in future studies (Borne et al., 2020; Lee et al., 2024, 2025; Lyu et al., 2021). The time-consuming manual definitions of primary, secondary, and PTS also limit the cortical expanse explored in each study, thus restricting the present study to LPC/LPOJ.”

Second, we utilized a young adult sample as this is what is the standard of the field when charting features of sulci for the first time (for example, Paus et al., 1996 Cerebral Cortex; Chiavaras & Petrides, 2000 Journal of Comparative Neurology; Segal & Petrides, 2012 European Journal of Neuroscience; Zlatkina & Petrides, 2014 Proceedings of the Royal Society B Biological Science; Sprung-Much & Petrides, 2018 Brain Structure & Function; Miller et al., 2021 The Journal of Neuroscience; Willbrand et al., 2022 Science Advances; Willbrand et al., 2023 Communications Biology; Drudik et al., 2023 Cerebral Cortex). Nevertheless, it is indeed crucial to confirm that this schematic is translatable to other age groups; however this exploration is beyond the scope of the present project and is for future investigation. We have added text to the Limitations subsection of the Discussion to emphasize the points (Pages 17-18):

“Additionally, the scope of the present study is limited in that the sample was only in young adults. This sample was selected as it is the standard of the field when charting features of sulci for the first time (for example, Paus et al. 1996; Chiavaras and Petrides 2000; Segal and Petrides 2012; Zlatkina and Petrides 2014; Sprung-Much and Petrides 2018; Miller et al. 2021; Willbrand et al. 2022; Willbrand et al. 2023; Drudik et al. 2023). Nevertheless, it is necessary to explore how well this updated schematic translates to different age groups, species, and clinical populations.”

Finally, it is worth mentioning that we have begun preliminary analyses on the translatability of this schematic, and have shown that it does hold in a pediatric sample (ages 6-18 years old; Author response image 1).

Example pediatric participant with all LPC/LOPJ sulci identified in both hemispheres.

Incidence rates for the variable pTS identified in the present work in a pediatric sample are included below (N = 79 participants)

(2) While the paper begins by describing four new sulci, only one is explored further in greater detail.

We focused on the slocs-v as it has a high incidence rate, making it amenable to our analytic pipelines relating sulci to cortical morphology, architecture, and function, as well as cognition (Miller et al., 2021 The Journal of Neuroscience; Voorhies et al., 2021 Nature Communications; Yao et al., 2022 Cerebral Cortex; Willbrand et al., 2022 Science Advances; Willbrand et al., 2023 The Journal of Neuroscience; Maboudian et al., 2024 The Journal of Neuroscience). However, we want to emphasize that throughout the paper there are multiple analyses that further describe the three more variable sulci: 1) detailing their sulcal patterning (Supplementary Tables 1-4) and 2) detailing their morphology and architecture (Supplementary Fig. 6). We do agree though that it is a worthwhile endeavor to further describe these sulci—especially if the data is readily available. As such, to complement our behavioral analysis identifying a relationship between the morphology of the consistent sulci and spatial orientation and considering the well-documented relationship between sulcal incidence and cognition (for review see Cachia et al., 2021 Frontiers in Neuroanatomy), we tested whether the number of variable sulci and the incidence of each variable sulcus specifically were related to spatial orientation. This procedure produced null results on all neuroanatomical variables, which we now mention in the Results (Page 11):

“Finally, as in prior work examining variably-present PTS in other cortical expanses (for example, (Amiez et al., 2018; Cachia et al., 2014; Fornito et al., 2004; Willbrand et al., 2024b), we assessed whether the presence/absence of the more variable PTS identified in the present work (slocs-d, pAngs-v, and pAngs-d) was related to spatial orientation, reasoning, and processing speed task performance. We identified no significant associations between the presence/absence of these sulci in either hemisphere with performance on these tests (ps > .05).”

(3) There is some tension between calling the discovered sulci new vs acknowledging they have already been reported, but not named.

To resolve this tension, we have revised the text to 1) ensure proper acknowledgment that sulci have been noticed in this region, 2) point out that these sulci were left unnamed and undescribed, and 3) emphasize that one of the primary goals of this project was to comprehensively detail the sulcal organization of this region at a precise, individual-level considering these often-overlooked sulci.

This is primarily done at the beginning of the Results (Pages 4-5), where we now write:

“Four previously undescribed small and shallow sulci in the lateral parieto-occipital junction (LPOJ)

In previous research in small sample sizes, neuroanatomists noticed shallow sulci in this cortical expanse, but did not describe them beyond including an unlabeled sulcus in their schematic at best (Supplementary Methods and Supplementary Figs. 1-4 for historical details). In the present study, we fully update this sulcal landscape considering these overlooked indentations. In addition to defining the 13 sulci previously described within the LPC/LPOJ, as well as the posterior superior temporal cortex in individual participants (Methods) (Petrides, 2019), we could also identify as many as four small and shallow PTS situated within the LPC/LPOJ that were highly variable across individuals and left undescribed until now (Supplementary Methods and Supplementary Figs. 1-4). Though we officially name and characterize features of these sulci in this paper for the first time, it is necessary to note that the location of these four sulci is consistent with the presence of variable “accessory sulci” in this cortical expanse mentioned in prior modern and classic studies (Supplementary Methods). For four example hemispheres with these 13-17 sulci identified, see Fig. 1a (Supplementary Fig. 5 for all hemispheres).”

(4) The anatomy of the sulci, as opposed to their relation to other sulci, could be described in greater detail.

To detail these sulci above and beyond their relation to other sulci, we document the anatomical metrics of all sulci in Supplemental Figure 6:

Results (Page 8):

The morphological and architectural features of all LPC/LPOJ sulci are described in Supplementary Fig. 6.

Changes from the earlier version to the revised version

Removed in the revised version: Abstract

Abstract

Recent work has uncovered relationships between evolutionarily new small and shallow cerebral indentations, or sulci, and human behavior. Yet, this relationship remains unexplored in the lateral parietal cortex (LPC) and the lateral parieto-occipital junction (LPOJ). After defining thousands of sulci in a young adult cohort, we revised the previous LPC/LPOJ sulcal landscape to include four previously overlooked, small, shallow, and variable sulci. One of these sulci (ventral supralateral occipital sulcus, slocs-v) is present in nearly every hemisphere and is morphologically, architecturally, and functionally dissociable from neighboring sulci. A data-driven, model-based approach, relating sulcal depth to behavior, further revealed that the morphology of only a subset of LPC/LPOJ sulci, including the slocs-v, is related to performance on a spatial orientation task. Our findings build on classic neuroanatomical theories and identify new neuroanatomical targets for future “precision imaging” studies exploring the relationship among brain structure, brain function, and cognitive abilities in individual participants.

Introduction

Earlier version: A fundamental goal in psychology and neuroscience is to understand the complex relationship between brain structure and brain function, as well as how that relationship provides a scaffold for efficient cognition and behavior. Of all the neuroanatomical features to target, recent work shows that morphological features of the shallower, later developing, hominoid-specific indentations of the cerebral cortex (also known as putative tertiary sulci, PTS) are not only functionally and cognitively meaningful, but also are particularly impacted by multiple brain-related disorders and aging (Amiez et al., 2018, 2019; Ammons et al., 2021; Cachia et al., 2021; Fornito et al., 2004; Garrison et al., 2015; Harper et al., 2022; Lopez-Persem et al., 2019; Maboudian et al., 2024; Miller et al., 2021; Nakamura et al., 2020; Parker et al., 2023; Ramos Benitez et al., 2024; Voorhies et al., 2021; Weiner, 2019; Willbrand, Ferrer, et al., 2023; Willbrand, Parker, et al., 2022; Willbrand, Voorhies, et al., 2022; Yao et al., 2022). The combination of these findings provides growing support for a classic theory proposing that the late gestational emergence of these PTS in gestation within association cortices, as well as their prolonged development, may co-occur with specific functional and microstructural features that could support specific cognitive abilities that also have a protracted development (Sanides, 1964). Nevertheless, despite the developmental, evolutionary, functional, cognitive, and theoretical relevance of these findings, PTS have mainly been restricted to only a subset of association cortices such as the prefrontal, cingulate, and ventral occipitotemporal cortices (Amiez et al., 2018, 2019; Ammons et al., 2021; Cachia et al., 2021; Fornito et al., 2004; Garrison et al., 2015; Harper et al., 2022; Hathaway et al., 2023; Lopez-Persem et al., 2019; Miller et al., 2020, 2021; Nakamura et al., 2020; Parker et al., 2023; Voorhies et al., 2021; Weiner, 2019; Willbrand, Ferrer, et al., 2023; Willbrand, Maboudian, et al., 2023; Willbrand, Parker, et al., 2022; Willbrand, Voorhies, et al., 2022; Yao et al., 2022). Thus, examining the relationship among these PTS relative to architectonic and functional features of the cerebral cortex, as well as relative to cognition, remains uncharted in other association cortices such as the lateral parietal cortex (LPC).

Revised version: A fundamental goal in psychology and neuroscience is to understand the complex relationship between brain structure and brain function, as well as how that relationship provides a scaffold for efficient cognition and behavior. Of all the neuroanatomical features to target, recent work shows that morphological features of the shallower, later developing, hominoid-specific indentations of the cerebral cortex (also known as putative tertiary sulci, PTS) are not only functionally and cognitively meaningful, but also are particularly impacted by multiple brain-related disorders and aging (Amiez et al., 2019, 2018; Ammons et al., 2021; Cachia et al., 2021; Fornito et al., 2004; Garrison et al., 2015; Harper et al., 2022; Lopez-Persem et al., 2019; Maboudian et al., 2024; Miller et al., 2021; Nakamura et al., 2020; Parker et al., 2023; Ramos Benitez et al., 2024; Voorhies et al., 2021; Weiner, 2019; Willbrand et al., 2023b, 2022a, 2022b; Yao et al., 2022). The combination of these findings provides growing support for a classic theory proposing that the late gestational emergence of these PTS in gestation within association cortices, as well as their prolonged development, may co-occur with specific functional and microstructural features that could support specific cognitive abilities that also have a protracted development (Sanides, 1964). Nevertheless, despite the developmental, evolutionary, functional, cognitive, and theoretical relevance of these findings, PTS have mainly been restricted to only a subset of association cortices such as the prefrontal, cingulate, and ventral occipitotemporal cortices (Amiez et al., 2019, 2018; Ammons et al., 2021; Cachia et al., 2021; Fornito et al., 2004; Garrison et al., 2015; Harper et al., 2022; Hathaway et al., 2023; Lopez-Persem et al., 2019; Miller et al., 2021, 2020; Nakamura et al., 2020; Parker et al., 2023; Voorhies et al., 2021; Weiner, 2019; Willbrand et al., 2023b, 2023c, 2022a, 2022b; Yao et al., 2022). Thus, examining the relationship among these PTS relative to architectonic and functional features of the cerebral cortex, as well as relative to cognition, remains uncharted in other association cortices such as the lateral parietal cortex (LPC).

Earlier version: As LPC is a cortical extent that has expanded extensively throughout evolution (Van Essen et al., 2018; Zilles et al., 2013), there is great interest in the structure and function of LPC in development, aging, across species, and in different patient populations. Yet, key gaps in knowledge relating individual differences in the structure of LPC to individual differences in the functional organization of LPC and cognitive performance remain for at least four main reasons. First, one line of recent work shows that LPC displays a much more complex sulcal patterning than previously thought (Drudik et al., 2023; Petrides, 2019; Segal & Petrides, 2012; Zlatkina & Petrides, 2014), while a second line of work shows that LPC is tiled with many maps and discrete functional regions spanning modalities and functions such as vision, memory, attention, action, haptics, and multisensory integration in addition to theory of mind, cognitive control, and subdivisions of the default mode network (Goodale & Milner, 1992; Harvey et al., 2013, 2015; Humphreys & Tibon, 2023; Konen & Kastner, 2008; Mackey et al., 2017; Schurz et al., 2017). Second, a majority of the time, the two lines of work are conducted independently from one another and the majority of human neuroimaging studies of LPC implement group analyses on average brain templates—which causes LPC sulci to disappear (Fig. 1). Third, despite the recently identified complexity of LPC sulcal patterning, recent studies have also uncovered previously overlooked PTS in association cortices (for example, in the posterior cingulate cortex; (Willbrand, Maboudian, et al., 2023; Willbrand, Parker, et al., 2022). Thus, fourth, it is unknown if additional LPC PTS are waiting to be detailed and if so, could improve our understanding of the structural-functional organization of LPC with potential cognitive insights as in other association cortices. Critically, while such findings would have developmental, evolutionary, functional, cognitive, and theoretical implications for addressing novel questions in future studies, they would also have translational applications as sulci serve as biomarkers in neurodevelopmental disorders (Ammons et al., 2021; Cachia et al., 2021; Garrison et al., 2015; Nakamura et al., 2020) and “corridors” for neurosurgery (Tomaiuolo et al., 2022; Tomaiuolo & Giordano, 2016).

Revised version: As LPC is a cortical extent that has expanded extensively throughout evolution (Van Essen et al., 2018; Zilles et al., 2013), there is great interest in the structure and function of LPC in development, aging, across species, and in different patient populations. Yet, key gaps in knowledge relating individual differences in the structure of LPC to individual differences in the functional organization of LPC and cognitive performance remain for at least four main reasons. First, one line of recent work shows that LPC displays a much more complex sulcal patterning than previously thought (Drudik et al., 2023; Petrides, 2019; Segal and Petrides, 2012; Zlatkina and Petrides, 2014), while a second line of work shows that LPC is tiled with many maps and discrete functional regions spanning modalities and functions such as vision, memory, attention, action, haptics, and multisensory integration in addition to theory of mind, cognitive control, and subdivisions of the default mode network (Goodale and Milner, 1992; Harvey et al., 2015, 2013; Humphreys and Tibon, 2023; Konen and Kastner, 2008; Mackey et al., 2017; Schurz et al., 2017). Second, a majority of the time, the two lines of work are conducted independently from one another and the majority of human neuroimaging studies of LPC implement group analyses on average brain templates—which causes LPC sulci to disappear (Fig. 1). Third, despite the recently identified complexity of LPC sulcal patterning, recent studies have also uncovered previously overlooked PTS in association cortices (for example, in the posterior cingulate cortex; Willbrand et al., 2023c, 2022a). Thus, fourth, it is unknown if additional LPC PTS are waiting to be detailed and if so, could improve our understanding of the structural-functional organization of LPC with potential cognitive insights as in other association cortices. Critically, while such findings would have developmental, evolutionary, functional, cognitive, and theoretical implications for addressing novel questions in future studies, they would also have translational applications as sulci serve as biomarkers in neurodevelopmental disorders (Ammons et al., 2021; Cachia et al., 2021; Garrison et al., 2015; Nakamura et al., 2020) and “corridors” for neurosurgery (Tomaiuolo and Giordano, 2016).

Four previously undefined small and shallow sulci in the lateral parieto-occipital junction (LPOJ).

Earlier version: a. Four example inflated (top) and pial (bottom) left hemisphere cortical surfaces displaying the 13-17 sulci manually identified in the present study. Each hemisphere contains 1-4 of the previousl undefined and variable LOC/LPOJ sulci (slocs and pAngs). Each sulcus is numbered according to the legend. b. Criteria for defining slocs and pAngs components. (i) Slocs-v is the cortical indentation between the cSTS3 and lTOS. (ii) Slocs-d is the indentation between cSTS3/cSTS2 and IPS-PO. (iii) pAngs-v is the indentation between the cSTS2 and pips. (iv) pAngs-d is the indentation between cSTS2/cSTS1 and IPS. c. The variability of the slocs and pAng components can cause them to disappear when individual surfaces are averaged together. Left to right: (i) 10 HCP participants, (ii) 20 HCP participants, (iii) 100 HCP participants, and iv) 650 HCP participants. The disappearance of these sulci on average surfaces, which are often used for group analyses in neuroimaging research, emphasizes the importance of defining these structures in individual hemispheres.

Revised version: a. Four example inflated (top) and pial (bottom) left hemisphere cortical surfaces displaying the 13-17 sulci manually identified in the present study. Each hemisphere contains 1–4 of the previousl undefined and variable LOC/LPOJ sulci (slocs and pAngs). Each sulcus is numbered according to the legend. b. Criteria for defining slocs and pAngs components. (i) Slocs-v is the cortical indentation between the cSTS3 and lTOS. (ii) Slocs-d is the indentation between cSTS3/cSTS2 and IPS-PO. (iii) pAngs-v is the indentation between the cSTS2 and pips. (iv) pAngs-d is the indentation between cSTS2/cSTS1 and IPS. c. The variability of the slocs and pAng components can cause them to disappear when individual surfaces are averaged together. Left to right: (i) 10 Human Connectome Project (HCP) participants, (ii) 20 HCP participants, (iii) 100 HCP participants, and iv) 650 HCP participants. The disappearance of these sulci on average surfaces, which are often used for group analyses in neuroimaging research, emphasizes the importance of defining these structures in individual hemispheres.

Earlier version: In the present study, we first manually defined LPC sulci in 144 young adult hemispheres using the most recent definitions of LPC sulci (Petrides, 2019). By manually labeling over 2,000 sulci, we detail four previously undescribed (Supplementary Methods and Supplementary Figs. 1-4 for historical details) sulci in the cortical expanse between the caudal branches of the superior temporal sulcus (cSTS) and two parts of the intraparietal sulcus (IPS)—a cortical expanse recently referenced as containing sensory “bridge” regions of the temporal-parietal-occipital junction (Glasser et al., 2016)—which we term the supralateral occipital sulci (ventral: slocs-v; dorsal: slocs-d) and posterior angular sulci (ventral: pAngs-d; dorsal: pAngs-d). We then utilized morphological (depth and surface area), architectural (gray matter thickness and myelination), and functional (resting-state functional connectivity) data available in each participant to assess whether the most common of these structures (slocs-v) was dissociable from surrounding sulci. Finally, we assessed whether the updated view of the LPC/LPOJ sulcal landscape provided cognitive insights using a model-based, data-driven approach (Voorhies et al., 2021) relating sulcal morphology to behavior on tasks known to activate regions within this cortical expanse [for example, reasoning and spatial orientation (Gur et al., 2000; Karnath, 1997; Vendetti & Bunge, 2014; Wendelken, 2014)].

Revised version: In the present study, we first manually defined LPC sulci in 144 young adult hemispheres using the most recent definitions of LPC sulci (Petrides, 2019). By manually labeling over 2,000 sulci, we detail four previously undescribed (Supplementary Methods and Supplementary Figs. 1–4 for historical details) sulci in the cortical expanse between the caudal branches of the superior temporal sulcus (cSTS) and two parts of the intraparietal sulcus (IPS)—a cortical expanse recently referenced as containing sensory “bridge” regions of the temporal-parietal-occipital junction (Glasser et al., 2016)—which we term the supralateral occipital sulci (ventral: slocs-v; dorsal: slocs-d) and posterior angular sulci (ventral: pAngs-d; dorsal: pAngs-d). We then utilized morphological (depth and surface area), architectural (gray matter thickness and myelination), and functional (resting-state functional connectivity) data available in each participant to assess whether the most common of these structures (slocs-v) was dissociable from surrounding sulci. Finally, we assessed whether the updated view of the LPC/LPOJ sulcal landscape provided cognitive insights using a model-based, data-driven approach (Voorhies et al., 2021) relating sulcal morphology to behavior on tasks known to activate regions within this cortical expanse (for example, reasoning and spatial orientation; Gur et al., 2000; Karnath, 1997; Vendetti and Bunge, 2014; Wendelken, 2014).

Results

Earlier version: Four previously undefined small and shallow sulci in the lateral parieto-occipital junction (LPOJ)

Revised version: Four previously undescribed small and shallow sulci in the lateral parieto-occipital junction (LPOJ)

Earlier version: In previous research in small sample sizes, neuroanatomists noticed shallow sulci in this cortical expanse (Supplementary Methods and Supplementary Figs. 1-4 for historical details). In the present study, we fully update this sulcal landscape considering these overlooked indentations. In addition to defining the 13 sulci previously described within the LPC/LPOJ, as well as the posterior superior temporal cortex (Methods) (Petrides, 2019) in individual participants, we could also identify as many as four small and shallow PTS situated within the LPC/LPOJ that were highly variable across individuals and uncharted until now (Supplementary Methods and Supplementary Figs. 1-4). Macroanatomically, we could identify two sulci between the cSTS3 and the IPS-PO/lTOS ventrally and two sulci between the cSTS2 and the pips/IPS dorsally. We focus our analyses on the slocs-v since it was identifiable in nearly every hemisphere.

Revised version: In previous research in small sample sizes, neuroanatomists noticed shallow sulci in this cortical expanse, but did not describe them beyond including an unlabeled sulcus in their figures and did not consider individual differences (Supplementary Methods and Supplementary Figs. 1–4 for historical details). In the present study, we fully update this sulcal landscape considering these overlooked indentations. In addition to defining the 13 sulci previously described within the LPC/LPOJ, as well as the posterior superior temporal cortex in individual participants (Methods; Petrides, 2019), we could also identify as many as four small and shallow PTS situated within the LPC/LPOJ that were highly variable across individuals and left undescribed until now (Supplementary Methods and Supplementary Figs. 1–4). Though we officially name and characterize features of these sulci in this paper for the first time, it is necessary to note that the location of these four sulci is consistent with the presence of variable “accessory sulci” in this cortical expanse mentioned in prior modern and classic studies (Supplementary Methods). For four example hemispheres with these 13-17 sulci identified, see Fig. 1a (Supplementary Fig. 5 for all hemispheres).

Earlier version: Ventrally, we refer to these sulci as ventral (slocs-v; sulcus 6 in Fig. 1) and dorsal (slocs-d; sulcus 7 in Fig. 1) components of the supralateral occipital sulcus. The slocs-v, located between the posterior cSTS3 and lTOS, was present in 98.6% of hemispheres (left hemisphere: N = 71/72; right hemisphere: N = 71/72; Fig. 1). Conversely, the more variable slocs-d, located between the cSTS3 and IPS-PO, was present 68.0% of the time (left hemisphere: N = 50/72; right hemisphere: N = 48/72; Fig. 1). Dorsally, we refer to the other newly described sulci as the ventral (pAngs-v; sulcus 8 in Fig. 1) and dorsal (pAngs-d; sulcus 9 in Fig. 1) components of the posterior angular sulcus. The pAng components were more rare than the slocs components. Specifically, pAngs-v, located between cSTS2 and pips, was identifiable 31.3% of the time (19 left and 26 right hemispheres; Fig. 1). Located between cSTS2 and the IPS, pAngs-d was identifiable only 13.2% of the time (8 left and 11 right hemispheres; Fig. 1). These incidence rates were significantly different (GLM, main effect of sulcus: χ2(3)L=L166.53, p < .0001; no hemispheric effects: ps > .68). The slocs-v was more common than the other three sulci (ps < .0001), slocs-d was more common than the pAngs components (ps < .0001), and pAngs-v was more common than pAngs-d (p = .002).

Revised version: Macroanatomically, we could identify two sulci between the cSTS3 and the IPS-PO/lTOS ventrally and two sulci between the cSTS2 and the pips/IPS dorsally. Ventrally, we refer to these sulci as ventral (slocs-v; sulcus 6 in Fig. 1) and dorsal (slocs-d; sulcus 7 in Fig. 1) components of the supralateral occipital sulcus (slocs). The slocs-v, located between the posterior cSTS3 and lTOS, was present in 98.6% of hemispheres (left hemisphere: N = 71/72; right hemisphere: N = 71/72; Fig. 1). Conversely, the more variable slocs-d, located between the cSTS3 and IPS-PO, was present 68.0% of the time (left hemisphere: N = 50/72; right hemisphere: N = 48/72; Fig. 1). Dorsally, we refer to the other newly described sulci as the ventral (pAngs-v; sulcus 8 in Fig. 1) and dorsal (pAngs-d; sulcus 9 in Fig. 1) components of the posterior angular sulcus (pAngs). The pAngs components were more rare than the slocs components. Specifically, pAngs-v, located between cSTS2 and pips, was identifiable 31.3% of the time (19 left and 26 right hemispheres; Fig. 1). Located between cSTS2 and the IPS, pAngs-d was identifiable only 13.2% of the time (8 left and 11 right hemispheres; Fig. 1). These incidence rates were significantly different (GLM, main effect of sulcus: χ2(3)L=L166.53, p < .0001; no hemispheric effects: ps > .68). The slocs-v was more common than the other three sulci (ps < .0001), slocs-d was more common than the pAngs components (ps < .0001), and pAngs-v was more common than pAngs-d (p = .002). We could also identify these sulci in post-mortem hemispheres (Supplementary Figs. 2, 3), ensuring that these sulci were not an artifact of the cortical reconstruction process.

Earlier version: Beyond characterizing the incidence of sulci, it is also common in the neuroanatomical literature to qualitatively characterize sulci on the basis of fractionation and intersection with surrounding sulci (termed “sulcal types”; for examples in other cortical expanses, see (Chiavaras & Petrides, 2000; Drudik et al., 2023; Miller et al., 2021; Paus et al., 1996; Weiner et al., 2014; Willbrand, Parker, et al., 2022). All four sulci most commonly did not intersect with other sulci (see Supplementary Tables 1-4 for a summary of the sulcal types of the slocs and pAngs dorsal and ventral components). The sulcal types were also highly comparable between hemispheres (rs > .99, ps < .001). Though we characterize these sulci in this paper for the first time, the location of these four sulci is consistent with the presence of variable “accessory sulci” in this cortical expanse mentioned in prior modern and classic studies (Supplementary Methods). We could also identify these sulci in post-mortem hemispheres (Supplementary Figs. 2, 3), ensuring that these sulci were not an artifact of the cortical reconstruction process.

Revised version: Beyond characterizing the incidence of sulci, it is also common in the neuroanatomical literature to qualitatively characterize sulci on the basis of fractionation and intersection with surrounding sulci (termed “sulcal types”; for examples in other cortical expanses, see Chiavaras and Petrides, 2000; Drudik et al., 2023; Miller et al., 2021; Paus et al., 1996; Weiner et al., 2014; Willbrand et al., 2022a). All four sulci most commonly did not intersect with other sulci (see Supplementary Tables 1–4 for a summary of the sulcal types of the slocs and pAngs dorsal and ventral components). The sulcal types were also highly comparable between hemispheres (rs > .99, ps < .001).

Earlier version: Given that sulcal incidence and patterning is also sometimes related to demographic features (Cachia et al., 2021; Leonard et al., 2009; Wei et al., 2017), subsequent GLMs relating the incidence and patterning of the three more variable sulci (slocs-d, pAngs-v, and pAngs-d) to demographic features (age and gender) revealed no associations for any sulcus (ps > .05). Finally, to help guide future research on these newly- and previously-classified LPC/LPOJ sulci, we generated probabilistic maps of each of these 17 sulci and share them with the field with the publication of this paper (Supplementary Fig. 6; Data availability).

Revised version: Given that sulcal incidence and patterning is also sometimes related to demographic features (Cachia et al., 2021; Leonard et al., 2009; Wei et al., 2017), subsequent GLMs relating the incidence and patterning of the three more variable sulci (slocs-d, pAngs-v, and pAngs-d) to demographic features (age and gender) revealed no associations for any sulcus (ps > .05). Finally, to help guide future research on these newly- and previously-classified LPC/LPOJ sulci, we generated probabilistic maps of each of these 17 sulci and share them with the field with the publication of this paper (Fig. 2; Data availability).

Added in the revised version: Maximum probability maps of the 17 sulci identified in the present study.

Added in the revised version: Maximum probability map (MPMs) for the 17 LPC/LPOJ sulci on the inflated fsaverage cortical surface (sulci: dark gray; gyri: light gray; cortical surfaces are not to scale) in the left (right surface; LH) and right (left surface; RH) hemispheres. To generate the MPMs, each label was transformed from each individual to the fsaverage surface. For each vertex, the proportion of participants for whom that vertex i labeled as the given sulcus (the warmer the color, the higher the overlap) was calculated. In the cases in which the vertices for each component overlapped, the sulcus with the highest overlap across participants was assigned to that vertex. For visual clarity, the MPMs were thresholded to 20% overlap acros participants. Sulci are numbered according to Fig. 1. These sulcal MPMs can be used to guide the definition of LPC/LPOJ sulci in future studies.

The slocs-v is morphologically, architecturally, and functionally dissociable from nearby sulci

Earlier version: Given that the slocs-v was present in the majority of participants (98.6% across hemispheres) and the other three sulci were far more variable (<70% of hemispheres), we focused our analyses on this stable sulcal feature of the LPOJ. To do so, we first tested whether the slocs-v was morphologically (depth and surface area) and architecturally (gray matter thickness and myelination) distinct from the two sulci surrounding it: the cSTS3 and lTOS (Fig. 1). These metrics are provided for all 17 sulci in Supplementary Figure 7. An rm-ANOVA (within-participant factors: sulcus, metric, and hemisphere for standardized metric units) revealed a sulcus x metric interaction (F(4, 276.19) = 179.15, η2 = 0.38, p < .001). Post hoc tests showed four main differences: (i) the slocs-v was shallower than cSTS3 (p < .001) but not lTOS (p = .60), (ii) the slocs-v was smaller than both the cSTS3 and lTOS (ps < .001), (iii) the slocs-v showed thicker gray matter than both the cSTS3 and lTOS (ps < .001), and iv) the slocs-v was less myelinated than both the cSTS and lTOS (ps < .001; Fig. 2a). There was also a sulcus x metric x hemisphere interaction (F(4.20, 289.81) = 4.16, η2 = 0.01, p = .002; hemispheric effects discussed in Supplementary Results).

Revised version: Given that the slocs-v was present in the majority of participants (98.6% across hemispheres) and the other three sulci were far more variable (<70% of hemispheres), we focused our analyses on this stable sulcal feature of the LPOJ. To do so, we first tested whether the slocs-v was morphologically (depth and surface area) and architecturally (gray matter thickness and myelination) distinct from the two sulci surrounding it: the cSTS3 and lTOS (Fig. 1). An rm-ANOVA (within-participant factors: sulcus, metric, and hemisphere for standardized metric units) revealed a sulcus x metric interaction (F(4, 276.19) = 179.15, η2 = 0.38, p < .001). Post hoc tests showed four main differences: (i) the slocs-v was shallower than cSTS3 (p < .001) but not lTOS (p = .60), (ii) the slocs-v was smaller than both the cSTS3 and lTOS (ps < .001), (iii) the slocs-v showed thicker gray matter than both the cSTS3 and lTOS (ps < .001), and iv) the slocs-v was less myelinated than both the cSTS and lTOS (ps < .001; Fig. 3a). There was also a sulcus x metric x hemisphere interaction (F(4.20, 289.81) = 4.16, η2 = 0.01, p = .002; hemispheric effects discussed in Supplementary Results). The morphological and architectural features of all LPC/LPOJ sulci are described in Supplementary Fig. 6.

The slocs-v is morphologically, architecturally, and functionally dissociable from nearby sulci.

a. Radial plot displaying the morphological (upper metrics: depth, surface area) and architectural (lower metrics: cortical thickness, myelination) features of the slocs-v (gray), cSTS3 (blue), and lTOS (green). Each dot and solid line represents the mean. The dashed lines indicate ± standard error. These features are colored by sulcus (legend). Metrics are in standardized units. b. Radial plot displaying the connectivity fingerprints of these three sulci: the Dice Coefficient overlap (values from 0-1) between each component and individual-level functional connectivity parcellations (Kong et al., 2019).

Earlier version: We then tested whether the slocs-v was also functionally distinct from the cSTS3 and lTOS by leveraging resting-state network parcellations for each individual participant to quantify “connectivity fingerprints” for each sulcus in each hemisphere of each participant (Methods) (Kong et al., 2019). An rm-ANOVA (within-participant factors: sulcus, network, and hemisphere for Dice coefficient overlap) revealed a sulcus x network interaction (F(32, 2144) = 80.18, η2 = 0.55, p < .001). Post hoc tests showed that this interaction was driven by four effects: (i) the cSTS3 overlapped more with the Default A subnetwork than both the slocs-v and lTOS (ps < .001), (ii) the slocs-v overlapped more with the Default C subnetwork than the lTOS (p < .001) and marginally than the cSTS3 (p = .077), (iii) the slocs-v overlapped more with the Dorsal Attention A subnetwork than both the cSTS3 and lTOS (ps < .001), and iv) the lTOS overlapped more with the Visual A and Visual B subnetworks than both the cSTS3 and slocs-v (ps < .004; Fig. 2b). There was also a sulcus x network x hemisphere interaction (F(32, 2144) = 3.99, η2 = 0.06, p < .001; hemispheric effects discussed in Supplementary Results). Together, these results indicate that the slocs-v is a morphologically, architecturally, and functionally distinct structure from its sulcal neighbors, and thus, deserves a distinct neuroanatomical definition.

Revised version: We then tested whether the slocs-v was also functionally distinct from the cSTS3 and lTOS by leveraging resting-state network parcellations for each individual participant to quantify “connectivity fingerprints” for each sulcus in each hemisphere of each participant (Methods; Kong et al., 2019). An rm-ANOVA (within-participant factors: sulcus, network, and hemisphere for Dice coefficient overlap) revealed a sulcus x network interaction (F(32, 2144) = 80.18, η2 = 0.55, p < .001). Post hoc tests showed that this interaction was driven by four effects: (i) the cSTS3 overlapped more with the Default A subnetwork than both the slocs-v and lTOS (ps < .001), (ii) the slocs-v overlapped more with the Default C subnetwork than the lTOS (p < .001) and marginally than the cSTS3 (p = .077), (iii) the slocs-v overlapped more with the Dorsal Attention A subnetwork than both the cSTS3 and lTOS (ps < .001), and iv) the lTOS overlapped more with the Visual A and Visual B subnetworks than both the cSTS3 and slocs-v (ps < .004; Fig. 3b). There was also a sulcus x network x hemisphere interaction (F(32, 2144) = 3.99, η2 = 0.06, p < .001; hemispheric effects discussed in Supplementary Results). Together, these results indicate that the slocs-v is a morphologically, architecturally, and functionally distinct structure from its sulcal neighbors, and thus, deserves a distinct neuroanatomical definition.

Earlier version: We further found that the three caudal STS rami (Petrides, 2019; Segal & Petrides, 2012) and intermediate parietal sulci (aipsJ and pips) (Petrides, 2019; Zlatkina & Petrides, 2014) are morphologically, architecturally, and functionally distinct structures for the first time (to our knowledge), which empirically supports their distinctions with separate sulcal labels (Supplementary Results and Supplementary Fig. 8).

Revised version: We further found that the three caudal STS rami (Petrides, 2019; Segal and Petrides, 2012) and intermediate parietal sulci (aipsJ and pips; Petrides, 2019; Zlatkina and Petrides, 2014) are morphologically, architecturally, and functionally distinct structures for the first time (to our knowledge), which empirically supports their distinctions with separate sulcal labels (Supplementary Results and Supplementary Fig. 7).

The morphology of LPC/LPOJ sulci, including the slocs-v, is related to cognitive performance

Earlier version: Finally, leveraging a data-driven approach of cross-validated LASSO feature selection, we sought to determine whether sulcal depth, a main defining feature of sulci, related to cognitive performance (Methods). To do so, we primarily focused on spatial orientation and reasoning given that these abilities recruit multiple subregions of lateral parietal and/or occipital cortices (Gur et al., 2000; Karnath, 1997; Vendetti & Bunge, 2014; Wendelken, 2014). As in prior work (Voorhies et al., 2021; Willbrand et al., 2024; Willbrand, Ferrer, et al., 2023; Willbrand, Voorhies, et al., 2022; Yao et al., 2022), we chose the model at the alpha that minimized MSEcv. Participants with a slocs-v in both hemispheres and all behavioral metrics were included (N = 69). Due to their rarity (being in less than 70% of hemispheres at most), we did not include the slocs-d or pAng components in this analysis.

Revised version: Finally, leveraging a data-driven approach of cross-validated LASSO feature selection, we sought to determine whether sulcal depth, a main defining feature of sulci, related to cognitive performance (Methods). To do so, we primarily focused on spatial orientation and reasoning given that these abilities recruit multiple subregions of lateral parietal and/or occipital cortices (Gur et al., 2000; Karnath, 1997; Vendetti and Bunge, 2014; Wendelken, 2014). As in prior work (Maboudian et al., 2024; Voorhies et al., 2021; Willbrand et al., 2023b; Yao et al., 2022), we chose the model at the alpha that minimized MSEcv. Participants with a slocs-v in both hemispheres and all behavioral metrics were included (N = 69). Due to their rarity (being in less than 70% of hemispheres at most), we did not include the slocs-d or pAng components in this analysis.

Earlier version: This method revealed an association between spatial orientation scores and normalized sulcal depth in the left hemisphere (MSEcv = 25.63, alpha = 0.05; Fig. 3a), but not in the right hemisphere (MSEcv = 26.41, alpha = 0.3). Further, we found that no LPC/LPOJ sulci were selected for reasoning in either hemisphere (right: alpha = 0.3, MSE = 24.01; left: alpha = 0.3, MSE = 24.01). Six left hemisphere LPC/LPOJ sulci were related to spatial orientation task performance (Fig. 3a, b). Four of these sulci were positioned ventrally: cSTS3 (β = −9.77), slocs-v (β = −3.36), lTOS (β = −4.91), and mTOS (β = −0.06), whereas two were positioned dorsally: pips (β = 5.02), and SPS (β = 4.30; Fig. 3a, b). Using LooCV to construct models that predict behavior, the LASSO-selected model explained variation in spatial orientation score (R2cv = 0.06, MSEcv = 23.99) above and beyond a model with all left hemisphere sulci (R2 < 0.01, MSEcv = 27.12). This model also showed a moderate correspondence (rs = 0.29, p = .01; Fig. 3c) between predicted and actual measured scores. We then tested for anatomical and behavioral specificity using the AIC, which revealed two primary findings. First, we found that the LASSO-selected sulcal depth model outperformed a model using the cortical thickness of the six LASSO-selected sulci (R2cv < .01, MSEcv = 26.02, AICcortical thickness – AICsulcal depth = 2.19). This model also showed task specificity as these sulci outperformed a model with processing speed (R2cv < .01, MSEcv = 254.65, AICprocessing speed – AICspatial orientation = 63.57). Thus, our data-driven model explains a significant amount of variance on a spatial orientation task and shows behavioral and morphological specificity.

Revised version: This method revealed an association between spatial orientation scores and normalized sulcal depth in the left hemisphere (MSEcv = 25.63, alpha = 0.05; Fig. 4a), but not in the right hemisphere (MSEcv = 26.41, alpha = 0.3). Further, we found that no LPC/LPOJ sulci were selected for reasoning in either hemisphere (right: alpha = 0.3, MSE = 24.01; left: alpha = 0.3, MSE = 24.01). Six left hemisphere LPC/LPOJ sulci were related to spatial orientation task performance (Fig. 4a, b). Four of these sulci were positioned ventrally: cSTS3 (β = −9.77), slocs-v (β = −3.36), lTOS (β = −4.91), and mTOS (β = −0.06), whereas two were positioned dorsally: pips (β = 5.02), and SPS (β = 4.30; Fig. 4a, b). Using LooCV to construct models that predict behavior, the LASSO-selected model explained variation in spatial orientation score (R2cv = 0.06, MSEcv = 23.99) above and beyond a model with all left hemisphere sulci (R2 < 0.01, MSEcv = 27.12). This model also showed a moderate correspondence (rs = 0.29, p = .01; Fig. 4c) between predicted and actual measured scores. We then tested for anatomical and behavioral specificity using the AIC, which revealed two primary findings. First, we found that the LASSO-selected sulcal depth model outperformed a model using the cortical thickness of the six LASSO-selected sulci (R2cv < .01, MSEcv = 26.02, AICcortical thickness – AICsulcal depth = 2.19). This model also showed task specificity as these sulci outperformed a model with processing speed (R2cv < .01, MSEcv = 254.65, AICprocessing speed – AICspatial orientation = 63.57). Thus, our data-driven model explains a significant amount of variance on a spatial orientation task and shows behavioral and morphological specificity.

The morphology of LPC/LPOJ sulci, including the slocs-v, is related to cognitive performance.

Earlier version: a. Beta-coefficients for each left hemisphere LPC/LPOJ sulcus at a range of shrinking parameter values [alpha (α)]. Highlighted gray bar indicates coefficients at the chosen α-level. Bottom: Cross-validated mean-squared error (MSECV) at each α level. By convention, we selected the α that minimized the MSECV (dotted line). b. Inflated left hemisphere cortical surface from an example participant highlighting the two groups of sulci—dorsal positive (orange) and ventral negative (green)— related to spatial orientation performance. c. Spearman’s correlation (rs) between the measured and the predicted spatial orientation scores from the LASSO-selected model is shown in a.

Revised version: a. Beta-coefficients for each left hemisphere LPC/LPOJ sulcus at a range of shrinking parameter values (alpha, α). The highlighted gray bar indicates coefficients at the chosen α-level. Bottom: Cross-validated mean-squared error (MSECV) at each α level. By convention, we selected the α that minimized the MSECV (dotted line). b. Inflated left hemisphere cortical surface from an example participant highlighting the two groups of sulci—dorsal positive (orange) and ventral negative (green)—related to spatial orientation performance. c. Spearman’s correlation (rs) between the measured and the predicted spatial orientation scores from the LASSO-selected model is shown in a.

Added in the revised version: Finally, as in prior work examining variably-present PTS in other cortical expanses (for example, Amiez et al., 2018; Cachia et al., 2014; Fornito et al., 2004; Willbrand et al., 2024b), we assessed whether the presence/absence of the more variable PTS identified in the present work (slocs-d, pAngs-v, and pAngs-d) was related to spatial orientation, reasoning, and processing speed task performance. We identified no significant associations between the presence/absence of these sulci in either hemisphere with performance on these tests (ps > .05).

Discussion

Overview

Earlier version: In the present study, we examined the relationship between LPC/LPOJ sulcal morphology, functional connectivity fingerprints, and cognition. We report five main findings. First, while manually defining sulci in LPC/LPOJ across 144 hemispheres, we uncovered four small and shallow sulci that are not included in present or classic neuroanatomy atlases or neuroimaging software packages. Second, we found that the most common of these structures (the slocs-v; identifiable 98.6% of the time) was morphologically, architecturally, and functionally differentiable from nearby sulci. Third, using a model-based, data-driven approach quantifying the relationship between sulcal morphology and cognition, we found a relationship between the depths of six LPC/LPOJ sulci and performance on a spatial orientation processing task. Fourth, the model identified distinct dorsal and ventral sulcal networks in LPC/LPOJ: ventral sulci had negative weights while dorsal sulci had positive weights (Fig. 3b). These findings are consistent with previous neuroimaging work from Gur et al. (Gur et al., 2000) who demonstrated separate functional activations in dorsal parietal and the more ventrally situated occipital-parietal cortices for the judgment of line orientation task used in the present study. Fifth, the model identified that the slocs-v is cognitively relevant, further indicating the importance of this neuroanatomical structure. In the sections below, we discuss (i) the slocs-v relative to modern functional and cytoarchitectonic parcellations in the LPC/LPOJ, as well as anatomical connectivity to other parts of the brain, (ii) underlying anatomical mechanisms relating sulcal morphology and behavior more broadly, and (iii) limitations of the present study. Implications for future studies are distributed throughout each section.

Revised version: In the present study, we examined the relationship between LPC/LPOJ sulcal morphology, functional connectivity fingerprints, and cognition. We report five main findings. First, while manually defining sulci in LPC/LPOJ across 144 hemispheres, we uncovered four small and shallow sulci that are not included in present or classic neuroanatomy atlases or neuroimaging software packages. Second, we found that the most common of these structures (the slocs-v; identifiable 98.6% of the time) was morphologically, architecturally, and functionally differentiable from nearby sulci. Third, using a model-based, data-driven approach quantifying the relationship between sulcal morphology and cognition, we found a relationship between the depths of six LPC/LPOJ sulci and performance on a spatial orientation processing task. Fourth, the model identified distinct dorsal and ventral sulcal networks in LPC/LPOJ: ventral sulci had negative weights while dorsal sulci had positive weights (Fig. 4b). These findings are consistent with previous neuroimaging work from Gur et al. (2000) who demonstrated separate functional activations in dorsal parietal and the more ventrally situated occipital-parietal cortices for the judgment of line orientation task used in the present study. Fifth, the model identified that the slocs-v is cognitively relevant, further indicating the importance of this neuroanatomical structure. In the sections below, we discuss (i) the slocs-v relative to modern functional and cytoarchitectonic parcellations in the LPC/LPOJ, as well as anatomical connectivity to other parts of the brain, (ii) underlying anatomical mechanisms relating sulcal morphology and behavior more broadly, and (iii) limitations of the present study. Implications for future studies are distributed throughout each section.

The slocs-v relative to modern functional and cytoarchitectonic parcellations in the LPC/LPOJ, as well as anatomical connectivity to other parts of the brain

Earlier version: To lay the foundation for future studies relating the newly-described slocs-v to different anatomical and functional organizational features of LPC/LPOJ, we situate probabilistic predictions of slocs-v relative to probabilistic cortical areas identified using multiple modalities. For example, when examining the correspondence between the slocs-v and modern multimodal (HCP-MMP) (Glasser et al., 2016) and observer-independent cytoarchitectural (Julich-Brain atlas) (Amunts et al., 2020) areas (Methods), the slocs-v is located within distinct areas. In particular, the slocs-v aligns with the multimodally- and cytoarchitecturally-defined area PGp bilaterally and cytoarchitecturally-defined hIP4 in the right hemisphere (Fig. 4). In classic neuroanatomical terms (Cunningham, 1892), this indicates that the slocs-v is a putative “axial sulcus” for these regions, which future work can assess with analyses in individual participants.

Revised version: To lay the foundation for future studies relating the newly-described slocs-v to different anatomical and functional organizational features of LPC/LPOJ, we situate probabilistic predictions of slocs-v relative to probabilistic cortical areas identified using multiple modalities. For example, when examining the correspondence between the slocs-v and modern multimodal (Human Connectome Project multimodal parcellation, HCP-MMP; Glasser et al., 2016) and observer-independent cytoarchitectural (Julich-Brain atlas; Amunts et al., 2020) areas (Methods), the slocs-v is located within distinct areas. In particular, the slocs-v aligns with the multimodally- and cytoarchitecturally-defined area PGp bilaterally and cytoarchitecturally-defined hIP4 in the right hemisphere (Fig. 5). In classic neuroanatomical terms (Cunningham, 1892), this indicates that the slocs-v is a putative “axial sulcus” for these regions, which future work can assess with analyses in individual participants.

The slocs-v relative to modern functional and cytoarchitectonic parcellations in LPC/LPOJ.

Earlier version: a. Top: Left (LH) and right (RH) hemispheres of the inflated fsaverage surface with two areas from the modern HCP multimodal parcellation (HCP-MMP; blue) (Glasser et al., 2016) relative to an MPM of the slocs-v (warm colors indicate areas with at least 20% overlap across participants; Supplementary Fig. 6). Bottom: Same as top, except for two observer-independent cytoarchitectonic regions from the Julich-Brain Atlas (Amunts et al., 2020). b. Overlap between the slocs-v and each area (Methods). Each dot and solid line represents the mean. The dashed lines indicate ± standard error (left: gray; right: white).

Revised version: a. Top: Left (LH) and right (RH) hemispheres of the inflated fsaverage surface with two areas from the modern HCP multimodal parcellation (HCP-MMP; blue; Glasser et al., 2016) relative to an MPM of the slocs-v (warm colors indicate areas with at least 20% overlap across participants; Fig. 2). Bottom: Same as top, except for two observer-independent cytoarchitectonic regions from the Julich-Brain Atlas (Amunts et al., 2020). b. Overlap between the slocs-v and each area (Methods). Each dot and solid line represents the mean. The dashed lines indicate ± standard error (left: gray; right: white).

Earlier version: Aside from recent multimodal and observer-independent cytoarchitectonic parcellations, an immediate question is: What is the relationship between the slocs-v and other functional regions at this junction between the occipital and parietal lobes, as well as potential anatomical connectivity? For example, there are over a dozen visual field maps in the cortical expanse spanning the TOS, IPS-PO, and the IPS proper (see (i), (ii), and (iii), respectively in Fig. 5a) (Mackey et al., 2017). When projecting probabilistic locations of retinotopic maps from over 50 individuals from Wang and colleagues (L. Wang et al., 2015) (Methods), the slocs-v is likely located outside of visual field maps extending into this cortical expanse (Fig. 5a). Nevertheless, when also projecting the map of the mean R2 metric from the HCP retinotopy dataset from 181 participants shared by Benson and colleagues (Benson et al., 2018) (Methods), the slocs-v is in a cortical expanse that explains a significant amount of variance (left hemisphere: R2 = 19.29, R2max = 41.73; right hemisphere: R2mean = 21.17, R2max = 44.23; Fig. 5b).

Revised version: Aside from recent multimodal and observer-independent cytoarchitectonic parcellations, an immediate question is: What is the relationship between the slocs-v and other functional regions at this junction between the occipital and parietal lobes, as well as potential anatomical connectivity? For example, there are over a dozen visual field maps in the cortical expanse spanning the TOS, IPS-PO, and the IPS proper (see (i), (ii), and (iii), respectively in Fig. 6a; Mackey et al., 2017). When projecting probabilistic locations of retinotopic maps from over 50 individuals from Wang and colleagues (Wang et al., 2015; Methods), the slocs-v is likel located outside of visual field maps extending into this cortical expanse (Fig. 6a). Nevertheless, when also projecting the map of the mean R2 metric from the HCP retinotopy dataset from 181 participants shared by Benson and colleagues (Benson et al., 2018; Methods), the slocs-v is in a cortical expanse that explains a significant amount of variance (left hemisphere: R2 = 19.29, R2max = 41.73; right hemisphere: R2mean = 21.17, R2max = 44.23; Fig. 6b).

The slocs-v relative to retinotopy.

Earlier version: a. Top: Left (LH) and right (RH) hemispheres of the inflated fsaverage surface displaying the probabilistic locations of retinotopic maps from over 50 individuals from Wang and colleagues (L. Wang et al., 2015) (black outlines). The predicted slocs-v location from the MPMs is overlaid in orange (as in Fig. 4). (i), (ii), and (iii) point out the retinotopic maps in the cortical expanse spanning the TOS, IPS-PO, and IPS, respectively. b. Same format as in a, but with a map of the mean R2 metric from the HCP retinotopy dataset (Benson et al., 2018) overlayed on the fsaverage surfaces (thresholded between R2 values of 10% and 90%). This metric measures how well the fMRI time series at each vertex is explained by a population receptive field (pRF) model. The mean and max R2 values for the slocs-v MPM in each hemisphere are included below each surface.

Revised version: a. Top: Left (LH) and right (RH) hemispheres of the inflated fsaverage surface displaying the probabilistic locations of retinotopic maps from over 50 individuals from Wang and colleagues (Wang et al., 2015); black outlines). The predicted slocs-v location from the MPM is overlaid in orange (as in Fig. 4). (i), (ii), and (iii) point out the retinotopic maps in the cortical expanse spanning the TOS, IPS-PO, and IPS, respectively. b. Same format as in a, but with a map of the mean R2 metric from the HCP retinotopy dataset (Benson et al., 2018) overlaid on the fsaverage surfaces (thresholded between R2 values of 10% and 90%). This metric measures how well the fMRI time series at each vertex is explained by a population receptive field (pRF) model. The mean and max R2 values for the slocs-v MPM in each hemisphere are included below each surface.

Earlier version: In terms of anatomical connectivity, as the slocs-v co-localizes with cytoarchitectonically defined PGp (Fig. 4) and previous studies have examined the anatomical connectivity of the probabilistically defined PGp, we can glean insight regarding the anatomical connectivity of slocs-v from these previous studies (Caspers et al., 2011; J. Wang et al., 2012). This prior work showed that PGp was anatomically connected to temporo-occipital regions, other regions in the temporal lobe, middle and superior frontal cortex, as well as the inferior frontal cortex and insula (Caspers et al., 2011; J. Wang et al., 2012). Furthermore, the slocs-v appears to lie at the junction of scene-perception and place-memory activity (a transition that also consistently co-localizes with the HCP-MMP area PGp) as identified by Steel and colleagues (Steel et al., 2021). Of course, the location of the slocs-v relative to multimodal, cytoarchitectonic, and retinotopic areas, as well as the anatomical connectivity of the slocs-v, would need to be examined in individual participants, but the present work makes clear predictions for future studies as fleshed out here. To conclude this section, as the multimodal area PGp (Fig. 4) was recently proposed as a “transitional area” by Glasser and colleagues (Glasser et al., 2016) (Supplementary Table 5), future studies can also further functionally and anatomically test the transitional properties of slocs-v.

Revised version: In terms of anatomical connectivity, as the slocs-v co-localizes with cytoarchitectonically defined PGp (Fig. 5) and previous studies have examined the anatomical connectivity of the probabilistically defined PGp, we can glean insight regarding the anatomical connectivity of slocs-v from these previous studies (Caspers et al., 2011; Wang et al., 2012). This prior work showed that PGp was anatomically connected to temporo-occipital regions, other regions in the temporal lobe, middle and superior frontal cortex, as well as the inferior frontal cortex and insula (Caspers et al., 2011; Wang et al., 2012). Furthermore, the slocs-v appears to lie at the junction of scene-perception and place-memory activity (a transition that also consistently co-localizes with the HCP-MMP area PGp) as identified by Steel and colleagues (2021). Of course, the location of the slocs-v relative to multimodal, cytoarchitectonic, and retinotopic areas, as well as the anatomical connectivity of the slocs-v, would need to be examined in individual participants, but the present work makes clear predictions for future studies as fleshed out here. To conclude this section, as the multimodal area PGp (Fig. 5) was recently proposed as a “transitional area” by Glasser and colleagues (2016; Supplementary Table 5), future studies can also further functionally and anatomically test the transitional properties of slocs-v.

Underlying anatomical mechanisms relating sulcal morphology and behavior

Earlier version: In this section, we discuss potential anatomical mechanisms contributing to the relationship between sulcal depth and behavior in two main ways. First, long-range white matter fibers have a gyral bias, while short-range white matter fibers have a sulcal bias in which some fibers project directly from the deepest points of a sulcus (Cottaar et al., 2021; Reveley et al., 2015; K. Schilling et al., 2018; K. G. Schilling et al., 2023; Van Essen et al., 2014). As such, recent work hypothesized a close link between sulcal depth and short-range white matter properties (Bodin et al., 2021; Pron et al., 2021; Voorhies et al., 2021; Willbrand, Ferrer, et al., 2023; Yao et al., 2022): deeper sulci would reflect even shorter short-range white matter fibers, which would result in faster communication between local, cortical regions and in turn, contribute to improved cognitive performance. This increased neural efficiency could underlie individual differences in cognitive performance. Ongoing work is testing this hypothesis which can be further explored in future studies incorporating anatomical, functional, and behavioral measures, as well as computational modeling.

Revised version: In this section, we discuss potential anatomical mechanisms contributing to the relationship between sulcal depth and behavior in two main ways. First, long-range white matter fibers have a gyral bias, while short-range white matter fibers have a sulcal bias in which some fibers project directly from the deepest points of a sulcus (Cottaar et al., 2021; Reveley et al., 2015; Schilling et al., 2018, 2023; Van Essen et al., 2014). As such, recent work hypothesized a close link between sulcal depth and short-range white matter properties (Bodin et al., 2021; Pron et al., 2021; Voorhies et al., 2021; Willbrand et al., 2023b; Yao et al., 2022): deeper sulci would reflect even shorter short-range white matter fibers, which would result in faster communication between local, cortical regions and in turn, contribute to improved cognitive performance. This increased neural efficiency could underlie individual differences in cognitive performance. Ongoing work is testing this hypothesis which can be further explored in future studies incorporating anatomical, functional, and behavioral measures, as well as computational modeling.

Earlier version: Second, our model-based approach identified separate dorsal and ventral sulcal networks in which deeper sulci dorsally and shallower sulci ventrally contributed to the most explained variance on the spatial orientation task. A similar finding was identified by our previous work in the lateral prefrontal cortex (Yao et al., 2022). These previous and present findings may be explained by the classic anatomical compensation theory, which proposes that the size and depth of a sulcus counterbalance those of the neighboring sulci (Armstrong et al., 1995; Connolly, 1950; Zilles et al., 2013). Thus, a larger, deeper sulcus would be surrounded by sulci that are smaller and shallower, rendering the overall degree of cortical folding within a given region approximately equal (Armstrong et al., 1995; Connolly, 1950; Zilles et al., 2013). Future work can incorporate underlying white matter architecture into the compensation theory, as well as a recent modification that proposed to also incorporate local morphological features such as the deepest sulcal point (e.g., sulcal pit or sulcal root (Régis et al., 2005)), which has recently been shown to be related to different functional features of the cerebral cortex (Bodin et al., 2018; Leroy et al., 2015; Natu et al., 2021). Altogether, these and recent findings begin to build a multimodal mechanistic neuroanatomical understanding underlying the complex relationship between sulcal depth and cognition relative to other anatomical features.

Revised version: Second, our model-based approach identified separate dorsal and ventral sulcal networks in which deeper sulci dorsally and shallower sulci ventrally contributed to the most explained variance on the spatial orientation task. A similar finding was identified by our previous work in the lateral prefrontal cortex (Yao et al., 2022). These previous and present findings may be explained by the classic anatomical compensation theory, which proposes that the size and depth of a sulcus counterbalance those of the neighboring sulci (Armstrong et al., 1995; Connolly, 1950; Zilles et al., 2013). Thus, a larger, deeper sulcus would be surrounded by sulci that are smaller and shallower, rendering the overall degree of cortical folding within a given region approximately equal (Armstrong et al., 1995; Connolly, 1950; Zilles et al., 2013). Future work can incorporate underlying white matter architecture into the compensation theory, as well as a recent modification that proposed to also incorporate local morphological features such as the deepest sulcal point (for example, sulcal pit or sulcal root; Régis et al., 2005), which has recently been shown to be related to different functional features of the cerebral cortex (Bodin et al., 2018; Leroy et al., 2015; Natu et al., 2021). Altogether, these and recent findings begin to build a multimodal mechanistic neuroanatomical understanding underlying the complex relationship between sulcal depth and cognition relative to other anatomical features.

Limitations

Earlier version: The main limitation of our study is that presently, the most accurate methodology to define sulci —especially the small, shallow, and variable PTS—requires researchers to manually trace each structure on the cortical surface reconstructions. This method is limited due to the individual variability of cortical sulcal patterning (Fig. 1, Supplementary Fig. 5), which makes it challenging to identify sulci without extensive experience and practice. However, we anticipate that our probabilistic maps will provide a starting point and hopefully, expedite the identification of these sulci in new participants. This should accelerate the process of subsequent studies confirming the accuracy of our updated schematic of LPC/LOPJ. This manual method is also arduous and time-consuming, which, on the one hand, limits the sample size in terms of number of participants, while on the other, results in thousands of precisely defined sulci. This push-pull relationship reflects a broader conversation in the human brain mapping and cognitive neuroscience fields between a balance of large N studies and “precision imaging” studies in individual participants (Allen et al., 2022; Gratton et al., 2022; Naselaris et al., 2021; Rosenberg & Finn, 2022). Though our sample size is comparable to other studies that produced reliable results relating sulcal morphology to brain function and cognition (e.g., (Cachia et al., 2021; Garrison et al., 2015; Lopez-Persem et al., 2019; Miller et al., 2021; Roell et al., 2021; Voorhies et al., 2021; Weiner, 2019; Willbrand, Parker, et al., 2022; Willbrand, Voorhies, et al., 2022; Yao et al., 2022), ongoing work that uses deep learning algorithms to automatically define sulci should result in much larger sample sizes in future studies (Borne et al., 2020; Lyu et al., 2021). The time-consuming manual definitions of primary, secondary, and PTS also limit the cortical expanse explored in each study, thus restricting the present study to LPC/LPOJ.

Revised version: The main limitation of our study is that presently, the most accurate methodology to define sulci —especially the small, shallow, and variable PTS—requires researchers to manually trace each structure on the cortical surface reconstructions. This method is limited due to the individual variability of cortical sulcal patterning (Fig. 1, Supplementary Fig. 5), which makes it challenging to identify sulci without extensive experience and practice. However, we anticipate that our probabilistic maps (Fig. 2) will provide a starting point and hopefully, expedite the identification of these sulci in new participants. This should accelerate the process of subsequent studies confirming the accuracy of our updated schematic of LPC/LOPJ. This manual method is also arduous and time-consuming, which, on the one hand, limits the sample size in terms of number of participants, while on the other, results in thousands of precisely defined sulci. This push-pull relationship reflects a broader conversation in the human brain mapping and cognitive neuroscience fields between a balance of large N studies and “precision imaging” studies in individual participants (Gratton et al., 2022; Naselaris et al., 2021; Rosenberg and Finn, 2022). Though our sample size is comparable to other studies that produced reliable results relating sulcal morphology to brain function and cognition (for example, Cachia et al., 2021; Garrison et al., 2015; Lopez-Persem et al., 2019; Miller et al., 2021; Roell et al., 2021; Voorhies et al., 2021; Weiner, 2019; Willbrand et al., 2022a, 2022b; Yao et al., 2022), ongoing work that uses deep learning algorithms to automatically define sulci should result in much larger sample sizes in future studies (Borne et al., 2020; Lee et al., 2024, 2025; Lyu et al., 2021). The time-consuming manual definitions of primary, secondary, and PTS also limit the cortical expanse explored in each study, thus restricting the present study to LPC/LPOJ.

Added in the revised version: Additionally, the scope of the present study is limited in that the sample was only in young adults. This sample was selected as it is the standard of the field when charting features of PTS for the first time (for example, Chiavaras and Petrides, 2000; Drudik et al., 2023; Miller et al., 2021; Paus et al., 1996; Segal and Petrides, 2012; Sprung-Much and Petrides, 2018; Willbrand et al., 2023c, 2022a; Zlatkina and Petrides, 2014). Nevertheless, it is necessary to explore how well this updated sulcal landscape translates to different age groups, species, and clinical populations.

Earlier version: It is also worth noting that the morphological-behavioral relationship identified in the present study explains a modest amount of variance; however, the more important aspect of our findings is that multiple sulci identified in our model-based approach are recently-characterized sulci in LPC/LOPJ identified by our group and others (Petrides, 2019), and thus, the relationship would have been overlooked or lost if these sulci were not identified. Finally, although we did not focus on the relationship between the other three PTS (slocs-d, pAngs-v, and pAngs-d) to anatomical and functional features of LPC and cognition, given that variability in sulcal incidence is cognitively (Amiez et al., 2018; Cachia et al., 2021; Garrison et al., 2015; Willbrand et al., 2024; Willbrand, Voorhies, et al., 2022), anatomically (Amiez et al., 2021; Vogt et al., 1995), functionally (Lopez-Persem et al., 2019), and translationally (Clark et al., 2010; Le Provost et al., 2003; Meredith et al., 2012; Nakamura et al., 2020; Yücel et al., 2002, 2003) relevant, future work can also examine the relationship between the more variable slocs-d, pAngs-v, and pAngs-d and these features.

Revised version: It is also worth noting that the morphological-behavioral relationship identified in the present study explains a modest amount of variance; however, the more important aspect of our findings is that multiple sulci identified in our model-based approach are recently-characterized sulci in LPC/LOPJ identified by our group and others (Petrides, 2019), and thus, the relationship would have been overlooked or lost if these sulci were not identified. Finally, although we did not focus on the relationship between the other three PTS (slocs-d, pAngs-v, and pAngs-d) to anatomical and functional features of LPC and cognition, given that variability in sulcal incidence is cognitively (for review see, Cachia et al., 2021), anatomically (Amiez et al., 2021; Vogt et al., 1995), functionally (Lopez-Persem et al., 2019), and translationally (Clark et al., 2010; Le Provost et al., 2003; Meredith et al., 2012; Nakamura et al., 2020; Yücel et al., 2003, 2002) relevant, future work can also examine the relationship between the more variable slocs-d, pAngs-v, and pAngs-d and these features.

Conclusion

In conclusion, we uncovered four previously-undefined sulci in LPC/LPOJ and quantitatively showed that the slocs-v is a stable sulcal landmark that is morphologically, architecturally, and functionally differentiable from surrounding sulci. We further used a data-driven, model-based approach relating sulcal morphology to behavior, which identified different relationships of ventral and dorsal LPC/LPOJ sulcal networks contributing to the perception of spatial orientation. The model identified the slocs-v, further indicating the importance of this newly-described neuroanatomical structure. Altogether, this work provides a scaffolding for future “precision imaging” studies interested in understanding how anatomical and functional features of LPC/LPOJ relate to cognitive performance at the individual level.

Methods

Participants

Earlier version: Data for the young adult human cohort analyzed in the present study were from the Human Connectome Project (HCP) database (https://www.humanconnectome.org/study/hcp-young-adult/overview). Here, we used 72 randomly-selected participants, balanced for gender (following the terminology of the HCP data dictionary), from the HCP database (50% female, 22-36 years old, and 90% right-handed; there was no effect of handedness on our behavioral tasks; Supplementary Methods) that were also analyzed in several prior studies (Hathaway et al., 2023; Miller et al., 2020, 2021; Willbrand, Ferrer, et al., 2023; Willbrand, Maboudian, et al., 2023; Willbrand, Parker, et al., 2022). HCP consortium data were previously acquired using protocols approved by the Washington University Institutional Review Board (Mapping the Human Connectome: Structure, Function, and Heritability; IRB # 201204036). Informed consent was obtained from all participants.

Revised version: Data for the young adult human cohort analyzed in the present study were from the Human Connectome Project (HCP) database (https://www.humanconnectome.org/study/hcp-young-adult/overview). Here, we used 72 randomly-selected participants, balanced for gender (following the terminology of the HCP data dictionary), from the HCP database (50% female, 22-36 years old, and 90% right-handed; there was no effect of handedness on our behavioral tasks; Supplementary Methods) that were also analyzed in several prior studies (Hathaway et al., 2023; Maboudian et al., 2024; Miller et al., 2021, 2020; Willbrand et al., 2024a, 2023b, 2023c, 2022a). HCP consortium data were previously acquired using protocols approved by the Washington University Institutional Review Board (Mapping the Human Connectome: Structure, Function, and Heritability; IRB # 201204036). Informed consent was obtained from all participants.

Neuroimaging data acquisition

Anatomical T1-weighted (T1-w) MRI scans (0.8 mm voxel resolution) were obtained in native space from the HCP database. Reconstructions of the cortical surfaces of each participant were generated using FreeSurfer (v6.0.0), a software package used for processing and analyzing human brain MRI images (surfer.nmr.mgh.harvard.edu; Dale et al., 1999; Fischl et al., 1999). All subsequent sulcal labeling and extraction of anatomical metrics were calculated from these native space reconstructions generated through the HCP’s version of the FreeSurfer pipeline (Glasser et al., 2013).

Behavioral data

Earlier version: In addition to structural and functional neuroimaging data, the HCP also includes a wide range of behavioral metrics from the NIH toolbox (Barch et al., 2013). To relate LPC/LPOJ sulcal morphology to behavior, we leveraged behavioral data related to spatial orientation (Variable Short Penn Line Orientation Test), relational reasoning (Penn Progressive Matrices Test), and processing speed (Pattern Completion Processing Speed Test; Supplementary Methods for task details). We selected these tasks as previous functional neuroimaging studies have shown the crucial role of LPC/LPOJ in relational reasoning and spatial orientation (Gur et al., 2000; Karnath, 1997; Vendetti & Bunge, 2014; Wendelken, 2014), while our previous work relating sulcal morphology to cognition uses processing speed performance as a control behavioral task (Voorhies et al., 2021; Willbrand, Voorhies, et al., 2022).

Revised version: In addition to structural and functional neuroimaging data, the HCP also includes a wide range of behavioral metrics from the NIH toolbox (Barch et al., 2013). To relate LPC/LPOJ sulcal morphology to behavior, we leveraged behavioral data related to spatial orientation (Variable Short Penn Line Orientation Test), relational reasoning (Penn Progressive Matrices Test), and processing speed (Pattern Completion Processing Speed Test; Supplementary Methods for task details). We selected these tasks as previous functional neuroimaging studies have shown the crucial role of LPC/LPOJ in relational reasoning and spatial orientation (Gur et al., 2000; Karnath, 1997; Vendetti and Bunge, 2014; Wendelken, 2014), while our previous work relating sulcal morphology to cognition uses processing speed performance as a control behavioral task (Voorhies et al., 2021; Willbrand et al., 2024a, 2022b).

Anatomical analyses

Manual labeling of LPC sulci

Earlier version: Sulci were manually defined in 72 participants (144 hemispheres) guided by the most recent atlas by Petrides (Petrides, 2019), as well as recent empirical studies (Drudik et al., 2023; Segal & Petrides, 2012; Zlatkina & Petrides, 2014), which together offer a comprehensive definition of cerebral sulcal patterns, including PTS. For a historical analysis of sulci in this cortical expanse, please refer to Segal & Petrides (Segal & Petrides, 2012) and Zlatkina & Petrides (Zlatkina & Petrides, 2014). Our cortical expanse of interest was bounded by the following sulci and gyri: (i) the postcentral sulcus (PoCS) served as the anterior boundary, (ii) the superior temporal sulcus (STS) served as the inferior boundary, (iii) the superior parietal lobule (SPL) served as the superior boundary, and (iv) the medial and lateral transverse occipital sulci (mTOS and lTOS) served as the posterior boundary. We also considered the following sulci within this cortical expanse: the three different branches of the caudal superior temporal sulcus (posterior to anterior: cSTS3, 2, 1), the supramarginal sulcus (SmgS), posterior intermediate parietal sulcus (pips), sulcus of Brissaud (sB), anterior intermediate parietal sulcus of Jensen (aipsJ), paroccipital intraparietal sulcus (IPS-PO), intraparietal sulcus (IPS), and the superior parietal sulcus (SPS). Of note, the IPS-PO is the portion of the IPS extending ventrally into the occipital lobe. The IPS-PO was first identified as the paroccipital sulcus by Wilder (1886). There is often an annectant gyrus separating the horizontal portion of the IPS proper from the IPS-PO (Roell et al., 2021; Zlatkina & Petrides, 2014).

Revised version: Sulci were manually defined in 72 participants (144 hemispheres) guided by the most recent atlas by Petrides (2019), as well as recent empirical studies (Drudik et al., 2023; Segal and Petrides, 2012; Zlatkina and Petrides, 2014), which together offer a comprehensive definition of cerebral sulcal patterns, including PTS. For a historical analysis of sulci in this cortical expanse, please refer to Segal & Petrides (2012) and Zlatkina & Petrides (2014). Our cortical expanse of interest was bounded by the following sulci and gyri: (i) the postcentral sulcus (PoCS) served as the anterior boundary, (ii) the superior temporal sulcus (STS) served as the inferior boundary, (iii) the superior parietal lobule (SPL) served as the superior boundary, and (iv) the medial and lateral transverse occipital sulci (mTOS and lTOS) served as the posterior boundary. We also considered the following sulci within this cortical expanse: the three different branches of the caudal superior temporal sulcus (posterior to anterior: cSTS3, 2, 1), the supramarginal sulcus (SmgS), posterior intermediate parietal sulcus (pips), sulcus of Brissaud (sB), anterior intermediate parietal sulcus of Jensen (aipsJ), paroccipital intraparietal sulcus (IPS-PO), intraparietal sulcus (IPS), and the superior parietal sulcus (SPS). Of note, the IPS-PO is the portion of the IPS extending ventrally into the occipital lobe. The IPS-PO was first identified as the paroccipital sulcus by Wilder (1886). There is often an annectant gyrus separating the horizontal portion of the IPS proper from the IPS-PO (Roell et al., 2021; Zlatkina and Petrides, 2014).

Additionally, we identified as many as four previously uncharted and variable LPC/LPOJ PTS for the first time: the supralateral occipital sulcus (slocs; composed of ventral (slocs-v) and dorsal (slocs-d) components) and the posterior angular sulcus (pAngs; composed of ventral (pAngs-v) and dorsal (pAngs-d) components). In the Supplementary Methods and Supplementary Figs. 1–4, we discuss the slocs and pAngs within the context of modern and historical sources.

Earlier version: For each participant in each hemisphere, the location of each sulcus was confirmed by first trained independent raters (Y.T. and T.G.), then confirmed by a trained expert (E.H.W.), and finally, finalized by a neuroanatomist (K.S.W.). All LPC sulci were then manually defined in FreeSurfer using tksurfer tools, as in previous work (Hathaway et al., 2023; Miller et al., 2020, 2021; Parker et al., 2023; Voorhies et al., 2021; Willbrand, Ferrer, et al., 2023; Willbrand, Maboudian, et al., 2023; Willbrand, Parker, et al., 2022; Willbrand, Voorhies, et al., 2022; Yao et al., 2022), from which morphological and anatomical features were extracted. We defined LPC/LPOJ sulci for each participant based on the most recent schematics of sulcal patterning by Petrides (2019) as well as pial, inflated, and smoothed white matter (smoothwm) FreeSurfer cortical surface reconstructions of each individual. In some cases, the precise start or end point of a sulcus can be difficult to determine on a surface (Borne et al., 2020); however, examining consensus across multiple surfaces allowed us to clearly determine each sulcal boundary in each individual. For four example hemispheres with these 13-17 sulci identified, see Fig. 1a (Supplementary Fig. 5 for all hemispheres). The specific criteria to identify the slocs and pAngs are outlined in Fig. 1b.

Revised version: As this is the first time the sulcal expanse of LPC/LOPJ was comprehensively charted with a focus on pTS, the location of each sulcus was confirmed through a three-tiered procedure for each participant in each hemisphere. First, trained independent raters (Y.T. and T.G.) identified sulci. Second, these definitions were checked by a trained expert (E.H.W.). Third, these labels were finalized by a neuroanatomist (K.S.W.). We emphasize that this procedure has produced reproducible results in our prior work across the cortex (Hastings et al., 2024; Maboudian et al., 2024; Miller et al., 2021, 2020; Parker et al., 2023; Ramos Benitez et al., 2024; Voorhies et al., 2021; Willbrand et al., 2024a, 2024b, 2023b, 2023c, 2022a, 2022b; Yao et al., 2022). All LPC sulci were then manually defined and saved as .label files in FreeSurfer using tksurfer tools, from which morphological and anatomical features were extracted. We defined LPC/LPOJ sulci for each participant based on the most recent schematics of sulcal patterning by Petrides (2019) as well as pial, inflated, and smoothed white matter (smoothwm) FreeSurfer cortical surface reconstructions of each individual. In some cases, the precise start or end point of a sulcus can be difficult to determine on a surface (Borne et al., 2020); however, examining consensus across multiple surfaces allowed us to clearly determine each sulcal boundary in each individual. For four example hemispheres with these 13-17 sulci identified, see Fig. 1a (Supplementary Fig. 5 for all hemispheres). The specific criteria to identify the slocs and pAngs are outlined in Fig. 1b.

To test whether the incidence rates of the slocs and pAngs components were statistically different, we implemented a binomial logistic regression GLM with sulcus (slocs-v, slocs-d, pAngs-v, and pAngs-d) and hemisphere (left and right), as well as their interaction, as predictors for sulcal presence (0: absent, 1: present). Additional GLMs were run relating the incidence of the more variable sulci (slocs-d, pAngs-v, and pAngs-d) to demographic features (gender and age) were also run. GLMs were carried out with the glm function from the built-in stats R package. ANOVA χ2 tests were applied to each GLM with the Anova function from the car R package, from which results were reported.

Probability maps

Earlier version: Sulcal probability maps were generated to show the vertices with the highest alignment across participants for a given sulcus. To create these maps, the label file for each sulcus was transformed from the individual to the fsaverage surface with the FreeSurfer mri_label2label command (https://surfer.nmr.mgh.harvard.edu/fswiki/mri_label2label). Once each label was transformed into this common template space, we calculated the proportion of participants for which each vertex was labeled as the given sulcus with custom Python code (Miller et al., 2021; Voorhies et al., 2021). For vertices with overlap between sulci, we employed a “winner-take-all” approach such that the sulcus with the highest overlap across participants was assigned to that vertex. Alongside the thresholded maps, we also provide constrained maps [maximum probability maps (MPMs)] at 20% participant overlap to increase interpretability (20% MPMs shown in Supplementary Fig. 6). To aid future studies interested in investigating LPC/LPOJ sulci, we share these maps with the field (Data availability).

Revised version: Sulcal probability maps were generated to show the vertices with the highest alignment across participants for a given sulcus. To create these maps, the label file for each sulcus was transformed from the individual to the fsaverage surface with the FreeSurfer mri_label2label command (https://surfer.nmr.mgh.harvard.edu/fswiki/mri_label2label). Once each label was transformed into this common template space, we calculated the proportion of participants for which each vertex was labeled as the given sulcus with custom Python code (Miller et al., 2021; Voorhies et al., 2021). For vertices with overlap between sulci, we employed a “winner-take-all” approach such that the sulcus with the highest overlap across participants was assigned to that vertex. Alongside the thresholded maps, we also provide constrained maps (maximum probability maps, MPMs) at 20% participant overlap to increase interpretability (20% MPMs shown in Fig. 2). To aid future studies interested in investigating LPC/LPOJ sulci, we share these maps with the field (Data availability).

Extracting and comparing the morphological and architectural features from sulcal labels

Earlier version: Morphologically, we compared sulcal depth and surface area across sulci, as these are two of the primary morphological features used to define and characterize sulci (Armstrong et al., 1995; Chi et al., 1977; Leroy et al., 2015; Lopez-Persem et al., 2019; Miller et al., 2020, 2021; Natu et al., 2021; Sanides, 1964; Voorhies et al., 2021; Weiner, 2019; Welker, 1990; Willbrand, Ferrer, et al., 2023; Willbrand, Parker, et al., 2022; Yao et al., 2022). As in our prior work (Voorhies et al., 2021; Yao et al., 2022), mean sulcal depth values (in standard FreeSurfer units) were computed in native space from the .sulc file generated in FreeSurfer (Dale et al., 1999) with custom Python code (Voorhies et al., 2021). Briefly, depth values are calculated based on how far removed a vertex is from what is referred to as a “mid-surface,” which is determined computationally so that the mean of the displacements around this “mid-surface” is zero. Thus, generally, gyri have negative values, while sulci have positive values. Each depth value was also normalized by the deepest point in the given hemisphere. Surface area (mm2) was calculated with the FreeSurfer mris_anatomical_stats function (https://surfer.nmr.mgh.harvard.edu/fswiki/mris_anatomical_stats). The morphological features of all LPC/LPOJ sulci are documented in Supplementary Fig. 7.

Revised version: Morphologically, we compared sulcal depth and surface area across sulci, as these are two of the primary morphological features used to define and characterize sulci (Armstrong et al., 1995; Chi et al., 1977; Leroy et al., 2015; Lopez-Persem et al., 2019; Miller et al., 2021, 2020; Natu et al., 2021; Sanides, 1964; Voorhies et al., 2021; Weiner, 2019; Welker, 1990; Willbrand et al., 2023b, 2022a; Yao et al., 2022). As in our prior work (Voorhies et al., 2021; Yao et al., 2022), mean sulcal depth values (in standard FreeSurfer units) were computed in native space from the .sulc file generated in FreeSurfer (Dale et al., 1999) with custom Python code (Voorhies et al., 2021). Briefly, depth values are calculated based on how far removed a vertex is from what is referred to as a “mid-surface,” which is determined computationally so that the mean of the displacements around this “mid-surface” is zero. Thus, generally, gyri have negative values, while sulci have positive values. Each depth value was also normalized by the deepest point in the given hemisphere. Surface area (mm2) was calculated with the FreeSurfer mris_anatomical_stats function (https://surfer.nmr.mgh.harvard.edu/fswiki/mris_anatomical_stats). The morphological features of all LPC/LPOJ sulci are documented in Supplementary Fig. 6.

Earlier version: Architecturally, we compared cortical thickness and myelination, as in our prior work in other cortical expanses (Miller et al., 2021; Voorhies et al., 2021; Willbrand, Ferrer, et al., 2023; Willbrand, Parker, et al., 2022). Mean gray matter cortical thickness (mm) was extracted using the FreeSurfer mris_anatomical_stats function. To quantify myelin content, we used the T1-w/T2-w maps for each hemisphere, an in vivo myelination proxy (Glasser & Van Essen, 2011). To generate the T1-w/T2-w maps, two T1-w and T2-w structural MR scans from each participant were registered together and averaged as part of the HCP processing pipeline (Glasser et al., 2013). The averaging helps to reduce motion-related effects or blurring. Additionally, and as described by Glasser and colleagues (Glasser et al., 2013), the T1-w/T2-w images were bias-corrected for distortion effects using field maps. We then extracted the average T1-w/T2-w ratio values across each vertex for each sulcus using custom Python code (Miller et al., 2021). The architectural features of all LPC/LPOJ sulci are documented in Supplementary Fig. 7.

Revised version: Architecturally, we compared cortical thickness and myelination, as in our prior work in other cortical expanses (Miller et al., 2021; Voorhies et al., 2021; Willbrand et al., 2023b, 2022a). Mean gray matter cortical thickness (mm) was extracted using the FreeSurfer mris_anatomical_stats function. To quantify myelin content, we used the T1-w/T2-w maps for each hemisphere, an in vivo myelination proxy (Glasser and Van Essen, 2011). To generate the T1-w/T2-w maps, two T1-w and T2-w structural MR scans from each participant were registered together and averaged as part of the HCP processing pipeline (Glasser et al., 2013). The averaging helps to reduce motion-related effects or blurring. Additionally, and as described by Glasser and colleagues (2013), the T1-w/T2-w images were bias-corrected for distortion effects using field maps. We then extracted the average T1-w/T2-w ratio values across each vertex for each sulcus using custom Python code (Miller et al., 2021). The architectural features of all LPC/LPOJ sulci are documented in Supplementary Fig. 6.

Earlier version: To assess whether these four metrics differed between the slocs-v and surrounding sulci (cSTS3 and lTOS), we ran a repeated measure analysis of variance (rm-ANOVA) with the within-participant effects of sulcus (slocs-v, cSTS3, and lTOS), metric (surface area, depth, cortical thickness, and myelination), and hemisphere (left and right). Rm-ANOVAs (including sphericity correction) were implemented with the aov_ez function from the afex R package. Effect sizes for the ANOVAs are reported with the partial eta-squared metric (η2). Post-hoc analyses were computed with the emmeans function from the emmeans R package (p-values corrected with Tukey’s method). We also repeated these analyses for the three cSTS components (Petrides, 2019; Segal & Petrides, 2012) and the two intermediate parietal sulcal components (ips: aipsJ and pips (Petrides, 2019; Zlatkina & Petrides, 2014); detailed in the Supplementary Results and Supplementary Fig. 8) as these components, to our knowledge, have not been quantitatively compared in previous work.

Revised version: To assess whether these four metrics differed between the slocs-v and surrounding sulci (cSTS3 and lTOS), we ran a repeated measure analysis of variance (rm-ANOVA) with the within-participant effects of sulcus (slocs-v, cSTS3, and lTOS), metric (surface area, depth, cortical thickness, and myelination), and hemisphere (left and right). Rm-ANOVAs (including sphericity correction) were implemented with the aov_ez function from the afex R package. Effect sizes for the ANOVAs are reported with the partial eta-squared metric (η2). Post-hoc analyses were computed with the emmeans function from the emmeans R package (p-values corrected with Tukey’s method). We also repeated these analyses for the three cSTS components (Petrides, 2019; Segal and Petrides, 2012) and the two intermediate parietal sulcal components (ips: aipsJ and pips; Petrides, 2019; Zlatkina and Petrides, 2014; detailed in the Supplementary Results and Supplementary Fig. 7) as these components, to our knowledge, have not been quantitatively compared in previous work.

Functional analyses

Earlier version: To determine if the slocs-v is functionally distinct from surrounding sulci, we generated functional connectivity profiles using recently developed analyses (Miller et al., 2021; Willbrand, Bunge, et al., 2023; Willbrand, Parker, et al., 2022). First, we used resting-state network parcellations for each individual participant from Kong and colleagues (Kong et al., 2019), who generated individual network definitions by applying a hierarchical Bayesian network algorithm to produce maps for each of the 17 networks in individual HCP participants. Importantly, this parcellation was conducted blind to both cortical folding and our sulcal definitions. Next, we resampled the network profiles for each participant onto the fsaverage cortical surface, and then to each native surface using CBIG tools (https://github.com/ThomasYeoLab/CBIG). We then calculated the spatial overlap between a sulcus and each of the 17 individual resting-state networks via the Dice coefficient (Equation 1):

Revised version: To determine if the slocs-v is functionally distinct from surrounding sulci, we generated functional connectivity profiles using recently developed analyses (Miller et al., 2021; Willbrand et al., 2023a, 2022a). First, we used resting-state network parcellations for each individual participant from Kong and colleagues (2019), who generated individual network definitions by applying a hierarchical Bayesian network algorithm to produce maps for each of the 17 networks in individual HCP participants. Importantly, this parcellation was conducted blind to both cortical folding and our sulcal definitions. Next, we resampled the network profiles for each participant onto the fsaverage cortical surface, and then to each native surface using CBIG tools (https://github.com/ThomasYeoLab/CBIG). We then calculated the spatial overlap between a sulcus and each of the 17 individual resting-state networks via the Dice coefficient (Equation 1): This process of calculating the overlap between each sulcus and the 17-network parcellation generated a “connectivity fingerprint” for each sulcus in each hemisphere of each participant. We then ran an rm-ANOVA with within-participant factors of sulcus (slocs-v, cSTS3, and lTOS), network (17 networks), and hemisphere (left and right) to determine if the network profiles (i.e., the Dice coefficient overlap with each network) of the slocs-v was differentiable from the surrounding sulci (i.e., cSTS3 and lTOS). Here we discuss effects related to networks that at least showed minor overlap with one sulcus (i.e., Dice ≥ .10). As in the prior analysis, we also repeated these analyses for the three cSTS components and the two intermediate parietal sulcal components (Supplementary Results and Supplementary Fig. 7).

Removed in the revised version: This process of calculating the overlap between each sulcus and the 17-network parcellation generated a “connectivity fingerprint” for each sulcus in each hemisphere of each participant. We then ran an rm-ANOVA with within-participant factors of sulcus (slocs-v, cSTS3, and lTOS), network (17 networks), and hemisphere (left and right) to determine if the network profiles (i.e., the Dice coefficient overlap with each network) of the slocs-v was differentiable from the surrounding sulci (i.e., cSTS3 and lTOS). Here we discuss effects related to networks that at least showed minor overlap with one sulcus (i.e., Dice ≥ .10). As in the prior analysis, we also repeated these analyses for the three cSTS components and the two intermediate parietal sulcal components (Supplementary Results and Supplementary Fig. 8).

Behavioral analyses

Model selection

Earlier version: The analysis relating sulcal morphology to spatial orientation and/or reasoning consisted of using a cross-validated (CV) least absolute shrinkage and selection operator (LASSO) regression to select the sulci that explained the most variance in the data and determined how much variance is explained by sulcal depth as a predictor of behavior, as implemented in our previous work (Maboudian et al., 2023; Voorhies et al., 2021; Willbrand, Ferrer, et al., 2023; Yao et al., 2022). A LASSO regression is well suited to address our question since it facilitates the model selection process and increases the generalizability of a model by providing a sparse solution that reduces coefficient values and decreases variance in the model without increasing bias (Heinze et al., 2018). Further, regularization is recommended in cases where there are many predictors (X > 10), as in this study, because this technique guards against overfitting and increases the likelihood that a model will generalize to other datasets. A LASSO performs L1 regularization by applying a penalty, or shrinking parameter (alpha, α), to the absolute magnitude of the coefficients. In this manner, low coefficients are set to zero and eliminated from the model. Therefore, LASSO affords data-driven variable selection that results in simplified models containing only the most predictive features, in this case, sulci predicting cognitive performance. This methodology improves model interpretability and prediction accuracy, as well as protects against overfitting, which improves generalizability (Ghojogh & Crowley, 2019; Heinze et al., 2018).

Revised version: The analysis relating sulcal morphology to spatial orientation and/or reasoning consisted of using a cross-validated (CV) least absolute shrinkage and selection operator (LASSO) regression to select the sulci that explained the most variance in the data and determined how much variance is explained by sulcal depth as a predictor of behavior, as implemented in our previous work (Maboudian et al., 2024; Voorhies et al., 2021; Willbrand et al., 2023b; Yao et al., 2022). A LASSO regression is well suited to address our question since it facilitates the model selection process and increases the generalizability of a model by providing a sparse solution that reduces coefficient values and decreases variance in the model without increasing bias (Heinze et al., 2018). Further, regularization is recommended in cases where there are many predictors (X > 10), as in this study, because this technique guards against overfitting and increases the likelihood that a model will generalize to other datasets. A LASSO performs L1 regularization by applying a penalty, or shrinking parameter (alpha, α), to the absolute magnitude of the coefficients. In this manner, low coefficients are set to zero and eliminated from the model. Therefore, LASSO affords data-driven variable selection that results in simplified models containing only the most predictive features, in this case, sulci predicting cognitive performance. This methodology improves model interpretability and prediction accuracy, as well as protects against overfitting, which improves generalizability (Ghojogh and Crowley, 2019; Heinze et al., 2018).

Earlier version: The depths of all LPC/LPOJ sulci were included as predictors in the LASSO regression model (Supplementary Methods for details on demographic control variables). We used nested CV to optimize the shrinking parameter for the LASSO regression. By convention (Heinze et al., 2018), we selected the model parameters that minimized the CV mean squared error (MSEcv). Optimization was performed with the GridSearchCV function from the SciKit-learn package in Python. This function allowed us to determine the model parameters minimizing the MSEcv by performing an exhaustive search across a range of α values. Nested CV was done as non-nested CV leads to biased performance (Cawley & Talbot, 2010; Vabalas et al., 2019).

Revised version: The depths of all LPC/LPOJ sulci were included as predictors in the LASSO regression model (Supplementary Methods for details on demographic control variables). We used nested CV to optimize the shrinking parameter for the LASSO regression. By convention (Heinze et al., 2018), we selected the model parameters that minimized the CV mean squared error (MSEcv). Optimization was performed with the GridSearchCV function from the SciKit-learn package in Python. This function allowed us to determine the model parameters minimizing the MSEcv by performing an exhaustive search across a range of α values. Nested CV was done as non-nested CV leads to biased performance (Cawley and Talbot, 2010; Vabalas et al., 2019).

To evaluate the performance of the model selected by the LASSO regression and verify the result of our feature selection, we used linear regression with leave-one-out CV (LooCV) to fit these selected models and to compare various models. Specifically, we measured the model performance for the relevant behavioral task using nested model comparison. With LooCV, we compared the LASSO-selected model with the predictors to a model with all left hemisphere sulci as predictors. All regression models were implemented with functions from the SciKit-learn Python package.

Assessing morphological and behavioral specificity

Earlier version: To assess whether our findings generalized to other anatomical features, we considered cortical thickness, which is consistently studied in cognitive neuroscience studies relating morphology to cognition (Dickerson et al., 2008; Gogtay et al., 2004; Voorhies et al., 2021; Willbrand, Ferrer, et al., 2023; Yao et al., 2022). To do so, we replaced sulcal depth with cortical thickness as the predictive metric in our LASSO-selected model. As with depth, the model was fit to the data with LooCV. To compare the thickness model to the depth model, we used the Akaike Information Criterion (AIC), which provides an estimate of in-sample prediction error and is suitable for non-nested model comparison. By comparing AIC scores, we are able to assess the relative performance of the two models. If the ΔAIC is > 2, it suggests an interpretable difference between models. If the ΔAIC is > 10, it suggests a strong difference between models, with the lower AIC value indicating the preferred model (Wagenmakers & Farrell, 2004). To also ascertain whether the relationship between LPC/LPOJ sulcal depth and cognition is specific to spatial orientation performance, or transferable to other general measures of cognitive processing, we investigated the generalizability of the sulcal-behavior relationship to another widely used measure of cognitive functioning: processing speed (Kail & Salthouse, 1994). Specifically, we used LooCV to predict processing speed instead of spatial orientation score. As with thickness, we compared the two models with the AIC.

Revised version: To assess whether our findings generalized to other anatomical features, we considered cortical thickness, which is consistently studied in cognitive neuroscience studies relating morphology to cognition (Dickerson et al., 2008; Gogtay et al., 2004; Maboudian et al., 2024; Voorhies et al., 2021; Willbrand et al., 2023b; Yao et al., 2022). To do so, we replaced sulcal depth with cortical thickness as the predictive metric in our LASSO-selected model. As with depth, the model was fit to the data with LooCV. To compare the thickness model to the depth model, we used the Akaike Information Criterion (AIC), which provides an estimate of in-sample prediction error and is suitable for non-nested model comparison. By comparing AIC scores, we are able to assess the relative performance of the two models. If the ΔAIC is > 2, it suggests an interpretable difference between models. If the ΔAIC is > 10, it suggests a strong difference between models, with the lower AIC value indicating the preferred model (Wagenmakers and Farrell, 2004). To also ascertain whether the relationship between LPC/LPOJ sulcal depth and cognition is specific to spatial orientation performance, or transferable to other general measures of cognitive processing, we investigated the generalizability of the sulcal-behavior relationship to another widely used measure of cognitive functioning: processing speed (Kail and Salthouse, 1994). Specifically, we used LooCV to predict processing speed instead of spatial orientation score. As with thickness, we compared the two models with the AIC.

Added in the revised version: Assessing the relationship between variable presence of the slocs-d, pAngs-v, and pAngs-d and behavior

Added in the revised version: To test the relationship between the more variable PTS identified in the present work and behavior, we implemented t-tests to assess the presence of the slocs-d, pAngs-v, and pAngs-d in both the left and right hemisphere to spatial orientation, reasoning, and processing speed task performance.

Situating the slocs-v within modern group-level cortical parcellations

To putatively relate the slocs-v to modern multimodal (HCP multimodal parcellation, HCP-MMP; Glasser et al., 2016) and cytoarchitectural (Julich-Brain atlas; Amunts et al., 2020) regions of the cerebral cortex located in fsaverage template space, we quantified the Dice coefficient overlap between the slocs-v of each participant (resampled to fsaverage space) and the individual regions of interest comprising the HCP-MMP and Julich-Brain parcellations.

Retinotopic response mapping of LPC/LPOJ sulci

To assess whether any of the LPC/LPOJ sulci related to retinotopic representations, we leveraged population receptive field mapping data (Benson et al., 2018). For each sulcal MPM (as the retinotopic data were only available in this template space), we extracted the mean R2 values (i.e., the percentage of variance in each vertex explained by the population receptive field model) for vertices that showed meaningful retinotopic responses across participants (thresholded at R2 > 10%; Mackey et al., 2017).

Removed in the revised version: Supporting information